.

# Homework 2: Arrays and Tables

Please complete this notebook by filling in the cells provided.   

Throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

Each of the qeustions has 5 points.

**Deadline:**

This assignment is due Sunday, May 12th at 23:59. 

```python
# Import datascience module first
from datascience import *
```

## 1. Creating Arrays


**Question 1.** Make an array called `weird_numbers` containing the following numbers (**in the given order**):

1. -2
2. the sine of 1.2
3. 3
4. 5 to the power of the cosine of 1.2

*Hint:* `sin` and `cos` are functions in the `math` module.

```python
# Our solution involved one extra line of code before creating
# weird_numbers.
...
weird_numbers = ...
weird_numbers
```

<!-- grade_id: cell-f82d5c408289c0b2, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
import numpy
assert type(weird_numbers) == numpy.ndarray
```

**Question 2.** Make an array called `book_title_words` containing the following three strings: "Eats", "Shoots", and "and Leaves" (**in the given order**).

```python
book_title_words = ...
book_title_words
```

<!-- grade_id: cell-8cbf64918be698c5, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(book_title_words) == numpy.ndarray
```

Strings have a method called `join`.  `join` takes one argument, an array of strings.  It returns a single string.  Specifically, the value of `a_string.join(an_array)` is a single string that's the [concatenation](https://en.wikipedia.org/wiki/Concatenation) ("putting together") of all the strings in `an_array`, **except** `a_string` is inserted in between each string.

**Question 3.** Use the array `book_title_words` and the method `join` to make two strings:

1. "Eats, Shoots, and Leaves" (call this one `with_commas`)
2. "Eats Shoots and Leaves" (call this one `without_commas`)

*Hint:* If you're not sure what `join` does, first try just calling, for example, `"foo".join(book_title_words)` .

```python
with_commas = ...
without_commas = ...

# These lines are provided just to print out your answers.
print('with_commas:', with_commas)
print('without_commas:', without_commas)
```

<!-- grade_id: cell-29c148fe6aa5e986, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(with_commas) == str
assert type(without_commas) == str
```

## 2. Indexing Arrays


These exercises give you practice accessing individual elements of arrays.  In Python (and in many programming languages), elements are accessed by *index*, so the first element is the element at index 0.  

**Question 1.** The cell below creates an array of some numbers.  Set `third_element` to the third element of `some_numbers` (i.e., -6).

```python
some_numbers = make_array(-1, -3, -6, -10, -15)

third_element = ...
third_element
```

<!-- grade_id: cell-dbbb8703c0f562de, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(third_element) == int
```

**Question 2.** The next cell creates a table that displays some information about the elements of `some_numbers` and their order.  Run the cell to see the partially-completed table, then fill in the missing information in the cell (the strings that are currently "???") to complete the table.

```python
elements_of_some_numbers = Table().with_columns(
    "English name for position", make_array("first", "second", "???", "???", "fifth"),
    "Element",                   some_numbers)
elements_of_some_numbers
```

<!-- grade_id: cell-4cc0e1b9675f5c2a, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(elements_of_some_numbers.column("English name for position")) == numpy.ndarray
```

**Question 3.** You'll sometimes want to find the *last* element of an array.  Suppose an array has 142 elements.  What is the index of its last element?

```python
index_of_last_element = ...
```

<!-- grade_id: cell-0b114f57143d94c8, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(index_of_last_element) == int
```

More often, you don't know the number of elements in an array, its *length*.  (For example, it might be a large dataset you found on the Internet.)  The function `len` takes a single argument, an array, and returns the `len`gth of that array (an integer).

**Question 4.** The cell below loads an array called `president_birth_years`.  The last element in that array is the most recent birth year of any deceased president. Assign that year to `most_recent_birth_year`.

```python
president_birth_years = Table.read_table("president_births.csv").column('Birth Year')

most_recent_birth_year = ...
most_recent_birth_year
```

<!-- grade_id: cell-695ad4366d62c8ec, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(most_recent_birth_year) == int
```

**Question 5.** Finally, assign `sum_of_birth_years` to the sum of the first, tenth, and last birth year in `president_birth_years`

```python
sum_of_birth_years = ...
```

<!-- grade_id: cell-df1e9274832027f6, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(sum_of_birth_years) == int
```

## 3. Basic Array Arithmetic


**Question 1.** Multiply the numbers 42, 4224, 42422424, and -250 by 157.  For this question, **don't** use arrays.

```python
first_product = ...
second_product = ...
third_product = ...
fourth_product = ...
print(first_product, second_product, third_product, fourth_product)
```

<!-- grade_id: cell-245ef21a100415fb, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(first_product) == int and type(second_product) == int and type(third_product) == int and type(fourth_product) == int
```

**Question 2.** Now, do the same calculation, but using an array called `numbers` and only a single multiplication (`*`) operator.  Store the 4 results in an array named `products`.

```python
numbers = ...
products = ...
products
```

<!-- grade_id: cell-ccfc2b537fc09f14, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(numbers) == numpy.ndarray and type(products) == numpy.ndarray
```

**Question 3.** We've loaded an array of temperatures in the next cell.  Each number is the highest temperature observed on a day at a climate observation station, mostly from the US.  Since they're from the US government agency [NOAA](noaa.gov), all the temperatures are in Fahrenheit.  Convert them all to Celsius by first subtracting 32 from them, then multiplying the results by $\frac{5}{9}$. Make sure to **ROUND** each result to the nearest integer using the `np.round` function.

```python
max_temperatures = Table.read_table("temperatures.csv").column("Daily Max Temperature")

celsius_max_temperatures = ...
celsius_max_temperatures
```

<!-- grade_id: cell-8b77b3e7ae52c9ec, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(celsius_max_temperatures) == numpy.ndarray
```

**Question 4.** The cell below loads all the *lowest* temperatures from each day (in Fahrenheit).  Compute the size of the daily temperature range for each day.  That is, first convert the minimum and maximum tempratures to the Celsius. Then, compute the difference between each daily maximum temperature and the corresponding daily minimum temperature.  **Give your answer in Celsius!** Make sure **NOT** to round your answer for this question!

```python
min_temperatures = Table.read_table("temperatures.csv").column("Daily Min Temperature")

celsius_temperature_ranges = ...
celsius_temperature_ranges
```

<!-- grade_id: cell-daaf40db4beb5281, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(celsius_temperature_ranges) == numpy.ndarray
```

## 4. World Population


The cell below loads a table of estimates of the world population for different years, starting in 1950. The estimates come from the [US Census Bureau website](https://www.census.gov/en.html).

```python
world = Table.read_table("world_population.csv").select('Year', 'Population')
world.show(4)
```

The name `population` is assigned to an array of population estimates.

```python
population = world.column(1)
population
```

In this question, you will apply some built-in Numpy functions to this array.

<img src="array_diff.png" style="width: 600px;"/>

The difference function `np.diff` subtracts each element in an array by the element that preceeds it. As a result, the length of the array `np.diff` returns will always be one less than the length of the input array.

**Question 1.** Very often in data science, we are interested understanding how values change with time. Use `np.diff` and `np.max` (or just `max`) to calculate the largest annual change in population between any two consecutive years.

```python
largest_population_change = ...
largest_population_change
```

<!-- grade_id: cell-4e2ff7a5aa815be7, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(largest_population_change) == numpy.int64 or type(largest_population_change) == int
```

## 5. Old Faithful


Old Faithful is a geyser in Yellowstone that erupts every 44 to 125 minutes (according to [Wikipedia](https://en.wikipedia.org/wiki/Old_Faithful)). People are [often told that the geyser erupts every hour](http://yellowstone.net/geysers/old-faithful/), but in fact the waiting time between eruptions is more variable. Let's take a look.

**Question 1.** The first line below assigns `waiting_times` to an array of 272 consecutive waiting times between eruptions, taken from a classic 1938 dataset. Assign the names `shortest`, `longest`, and `average` so that the `print` statement is correct.

```python
waiting_times = Table.read_table('old_faithful.csv').column('waiting')

shortest = ...
longest = ...
average = ...

print("Old Faithful erupts every", shortest, "to", longest, "minutes and every", average, "minutes on average.")
```

<!-- grade_id: cell-9ac6ae778a19e6ce, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert 40 <= shortest <= 150
assert 40 <= longest <= 150
assert 40 <= average <= 150
```

**Question 2.** Assign `biggest_change` to the biggest change in waiting time between two consecutive eruptions. For example, the third eruption occurred after 74 minutes and the fourth after 62 minutes, so the change in waiting time was 74 - 62 = 12 minutes. 

*Hint*: The biggest change could be negative, but in the end, we want to return the absolute value of the biggest change.

```python
biggest_change = ...
biggest_change
```

<!-- grade_id: cell-2526dd457855c427, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert 0 <= biggest_change <= (125 - 44)
```

## 6. Tables


**Question 1.** Suppose you have 4 apples, 3 oranges, and 3 pineapples.  (Perhaps you're using Python to solve a high school Algebra problem.)  Create a table that contains this information.  It should have two columns: "fruit name" and "count".  Give it the name `fruits`.

**Note:** Use lower-case and singular words for the name of each fruit, like `"apple"`. Furthermore, keep the order of rows and columns in your table as exactly mentioned here. 

```python
fruits = ...
fruits
```

<!-- grade_id: cell-c62216460a13c607, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(fruits) == tables.Table
assert "apple" in fruits.column("fruit name")
assert "orange" in fruits.column("fruit name")
assert "pineapple" in fruits.column("fruit name")
```

**Question 2.** The file `inventory.csv` contains information about the inventory at a fruit stand.  Each row represents the contents of one box of fruit.  Load it as a table named `inventory`.

```python
inventory = ...
inventory
```

<!-- grade_id: cell-89719cf1e84cd0e8, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(inventory) == tables.Table
```

**Question 3.** The file `sales.csv` contains the number of fruit sold from each box last Saturday.  It has an extra column called "price per fruit (\$)" that's the price *per item of fruit* for fruit in that box.  The rows are in the same order as the `inventory` table.  Load these data into a table called `sales`.

```python
sales = ...
sales
```

<!-- grade_id: cell-54df2161f9bc8d88, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(sales) == tables.Table
```

**Question 4.** How many fruits did the store sell in total on that day?

```python
total_fruits_sold = ...
total_fruits_sold
```

<!-- grade_id: cell-c64affce8c8afa0b, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(total_fruits_sold) == numpy.int64 or type(total_fruits_sold) == int
```

**Question 5.** What was the store's total revenue (the total price of all fruits sold) on that day?

*Hint:* If you're stuck, think first about how you would compute the total revenue from just the grape sales.

```python
total_revenue = ...
total_revenue
```

<!-- grade_id: cell-39aa149acc27c2ea, points: 5, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert type(total_revenue) == numpy.float64 or type(total_revenue) == float
```