Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE".

# Homework 1: Causality and Expressions
Please complete this notebook by filling in the cells provided.   

Throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. 

You should start early so that you have time to get help if you're stuck.

**Deadline:**

This assignment is due Sunday, April 28th at 23:59. 

**Recommended Reading:**
- [What is Data Science](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

## 1. Characters in Little Women


We counted the number of times that the literary characters were named in each chapter of the classic book, [Little Women](https://www.inferentialthinking.com/chapters/01/3/1/literary-characters). In computer science, the word "character" also refers to a letter, digit, space, or punctuation mark; any single element of a text. The following code generates a scatter plot in which each dot corresponds to a chapter of Little Women. The horizontal position of a dot measures the number of periods (i.e., dots ".") in the chapter. The vertical position measures the total number of characters.

```python
# This cell contains code that hasn't yet been covered in the course,
# but you should be able to interpret the scatter plot it generates.

from datascience import *
from urllib.request import urlopen
import numpy as np
% matplotlib inline

little_women_url = 'https://www.inferentialthinking.com/data/little_women.txt'
chapters = urlopen(little_women_url).read().decode().split('CHAPTER ')[1:]
text = Table().with_column('Chapters', chapters)
Table().with_columns(
    'Periods',    np.char.count(chapters, '.'),
    'Characters', text.apply(len, 0)
    ).scatter(0)
```

**Question 1.** Around how many periods are there in the chapter with the most characters? Assign either 1, 2, 3, 4, or 5 to the name `characters_q1` below. (10 points)

1. 250
2. 390
3. 440
4. 32,000
5. 40,000

```python
characters_q1 = ...
```

<!-- grade_id: cell-70b2a735d6e614bb, points: 10, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert 1 <= characters_q1 <= 5
```

**Question 2.** Which of the following chapters has the most characters per period? Assign either 1, 2, or 3 to the name `characters_q2` below. (10 points)
1. The chapter with about 60 periods
2. The chapter with about 350 periods
3. The chapter with about 440 periods

```python
characters_q2 = ...
```

<!-- grade_id: cell-ec1fa7b7c7597c65, points: 10, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert 1 <= characters_q2 <= 3
```

## 2. Names and Assignment Statements


**Question 1.** When you run the following cell, Python produces a cryptic error message.

```python
4 = 2 + 2
```

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `names_q1` below to indicate your answer. (10 points)

1. Python is smart and already knows `4 = 2 + 2`.

2. `4` is already a defined number, and it doesn't make sense to make a number be a name for something else. In Python, "`x = 2 + 2`" means "assign `x` as the name for the value of `2 + 2`."

3. It should be `2 + 2 = 4`.

4. I don't get an error message. This is a trick question.

```python
names_q1 = ...
```

<!-- grade_id: cell-f41bd230d562438a, points: 10, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert 1 <= names_q1 <= 4
```

**Question 2.** When you run the following cell, Python will produce another cryptic error message.

```python
two = 3
six = two plus two
```

Choose the best explanation of what's wrong with the code and assign 1, 2, 3, or 4 to `names_q2` below to indicate your answer. (10 points)

1. The `plus` operation only applies to numbers, not the word "two".

2. The name "two" cannot be assigned to the number 3.

3. Two plus two is four, not six.

4. The name `two` cannot be followed directly by another name.

```python
names_q2 = ...
```

<!-- grade_id: cell-eee209aed1b5745a, points: 10, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert 1 <= names_q2 <= 4
```

**Question 3.** When you run the following cell, Python will, yet again, produce another cryptic error message.

```python
x = print(5)
y = x + 2
```

Choose the best explanation of what's wrong with the code and assign 1, 2, or 3 to `names_q3` below to indicate your answer. (10 points)

1. Python doesn't want `y` to be assigned.

2. The `print` operation is meant for displaying values to the programmer, not for assigning values!

3. What error message?

```python
names_q3 = ...
```

<!-- grade_id: cell-53a9cb7dec0a61dd, points: 10, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert 1 <= names_q3 <= 3
```

## 3. Job Opportunities & Education in Rural India


A [study](http://www.nber.org/papers/w16021.pdf) at UCLA investigated factors that might result in greater attention to the health and education of girls in rural India. One such factor is information about job opportunities for women. The idea is that if people know that educated women can get good jobs, they might take more care of the health and education of girls in their families, as an investment in the girls’ future potential as earners. Without the knowledge of job opportunities, the author hypothesizes that families invest less in their women, and instead, invest in their men. 

The study focused on 160 villages outside the capital of India, all with little access to information about call centers and similar organizations that offer job opportunities to women. In 80 of the villages chosen at random, recruiters visited the village, described the opportunities, recruited women who had some English language proficiency and experience with computers, and provided ongoing support free of charge for three years. In the other 80 villages, no recruiters visited and no other intervention was made.

At the end of the study period, the researchers recorded data about the school attendance and health of the children in the villages.

**Question 1.** Which statement best describes the *treatment* and *control* groups for this study? Assign either 1, 2, or 3 to the name `jobs_q1` below. (10 points)

1. The treatment group was the 80 villages visited by recruiters, and the control group was the other 80 villages with no intervention.

2. The treatment group was the 160 villages selected, and the control group was the rest of the villages outside the capital of India.

3. There is no clear notion of *treatment* and *control* group in this study.

```python
jobs_q1 = ...
```

<!-- grade_id: cell-56ac960141b9122e, points: 10, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert 1 <= jobs_q1 <= 3
```

**Question 2.** Was this an observational study or a randomized controlled experiment? Assign either 1, 2, or 3 to the name `jobs_q2` below. (10 points)

1. This was an observational study.

2. This was a randomized controlled experiment.  

3. This was a randomized observational study.

```python
jobs_q2 = ...
```

<!-- grade_id: cell-f66fd68bfcf91c2e, points: 10, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert 1 <= jobs_q2 <= 3
```

## 4. Differences between Universities


**Question 1.** Suppose you'd like to *quantify* how *dissimilar* two universities are, using three quantitative characteristics.  The US Department of Education data on [UW](https://collegescorecard.ed.gov/school/?236948-University-of-Washington-Seattle-Campus) and [Cal](https://collegescorecard.ed.gov/school/?110635-University-of-California-Berkeley) describes the following three traits (among many others):

| Trait                                | UW     | Cal    |
|--------------------------------------|--------|--------|
| Average annual cost to attend ($)    | 13,566 | 13,707 |
| Graduation rate (percentage)         | 83     | 91     |
| Socioeconomic Diversity (percentage) | 25     | 31     |

You decide to define the dissimilarity between two universities as the maximum of the absolute values of the 3 differences in their respective trait values.

Using this method, compute the dissimilarity between UW and CAL.  Name the result `dissimilarity`.  Use a single expression (a single line of code) to compute the answer.  Let Python perform all the arithmetic (like subtracting 91 from 83) rather than simplifying the expression yourself. The built-in `abs` function takes absolute values. (20 points)

```python
dissimilarity = ...
dissimilarity
```

<!-- grade_id: cell-bcd0316de4051e07, points: 20, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert isinstance(dissimilarity, int)
```

## 5. Studying the Survivors


The Reverend Henry Whitehead was skeptical of John Snow’s conclusion about the Broad Street pump. After the Broad Street cholera epidemic ended, Whitehead set about trying to prove Snow wrong.  (The history of the event is detailed [here](http://www.ncbi.nlm.nih.gov/pmc/articles/PMC1034367/pdf/medhist00183-0026.pdf).)

He realized that Snow had focused his analysis almost entirely on those who had died. Whitehead, therefore, investigated the drinking habits of people in the Broad Street area who had not died in the outbreak.

What is the main reason it was important to study this group? (10 points)

1. If Whitehead had found that many people had drunk water from the Broad Street pump and not caught cholera, that would have been evidence against Snow's hypothesis.

2. Survivors could provide additional information about what else could have caused the cholera, potentially unearthing another cause.

3. Through considering the survivors, Whitehead could have identified a cure for cholera.

**Note:** Whitehead ended up finding further proof that the Broad Street pump played the central role in spreading the disease to the people who lived near it. Eventually, he became one of Snow’s greatest defenders.

```python
# Assign survivor_answer to 1, 2, or 3
survivor_answer = ...
```

<!-- grade_id: cell-54676b354e431b05, points: 10, grade: true, solution: false -->
```python
# This assert command will generate validation errors if you do not assign a proper value to the variable.
assert 1 <= survivor_answer <= 3
```